Question 1: Write a sum function that takes two numbers and returns the sum of them.

<!-- grade_id: cell-b348391a5ef6f302, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1d60428185a3d3d1, points: 10, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(sum(5,4), 9)
assert_equal(sum(1,1020), 1021)
assert_equal(sum(-5,4), -1)
assert_equal(sum(-4,-10), -14)
```

```python

```